---
title: 进阶
description: ""
---

## Insert exercise title here

```yaml
type: TabExercise
key: 4e24eb3cd9
xp: 100
```



`@pre_exercise_code`
```{python}

```

***

```yaml
type: NormalExercise
key: 7581597928
xp: 50
```

`@instructions`
声明一个变量 savings

`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: 5e1d863f8d
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
