---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

请编写一段代码，实现打印 "你好,Python"
> exercise.

`@instructions`
请使用 `print` 函数，来打印内容。

在 print 函数中，放置要打印的具体内容即可，即：`你好，Python`

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 37a31f1869
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`
